---
title: Test
description: Test


---
## <<<New Exercise>>>

```yaml
type: VideoExercise
key: cfed7d04f3
lang: python
xp: 50
skills: 2
video_link: player.vimeo.com/video/154783078
```



---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: 60c88349a5
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: 24bd5c418e
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: 49c01fe990
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
---
## Sample exercise

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: b18e71f0ff
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
